# Week 2 Problem Set

## Homeworks

```python
%load_ext nb_mypy
%nb_mypy On
```

Output:
```
Version 1.0.5
```

**HW1.** *Min-Heap:* Write the following function to implement *min-heap*. A *min-heap* is a binary heap that satisfies the *min-heap property*. This property can be described as:

    For all nodes except the root:
    
    A[left(i)] >= A[i]
    A[right(i)] >= A[i]

- `min_child(heap, index)`: which returns the index of the node's smallest child. The node you are referring has index of value `index`
- `min_heapify(array, index, size)`: which moves the node at `index` down the tree so as to satisfy the *min-heap property*. The argument `index` is the index of the node you want to start moving down in the array. The argument `size` is the size of the heap. This size may be the same or less than the number of elements in the array. Hint: You may need the `min_child()` function.
- `build_min_heap(array)`: which build a *min-heap* from an arbitrary array of integers. This function should make use of `min_heapify(array, index)`. 

```python
# Copy over the implementations of left_of & right_of from the Cohort qns
def left_of(index: int) -> int:
    return index*2+1

def right_of(index: int) -> int:
    return index*2+2
```

```python
def min_child(heap: list[int|float], index: int, heap_size: int) -> int:
    lf, rg = left_of(index), right_of(index)
    if (lf >= heap_size):
        return -1
    if (rg < heap_size and heap[rg] < heap[lf]):
        return rg
    return lf
```

<!-- grade_id: hw11, points: 1, grade: true, solution: false -->
```python
minheap: list[int|float] = [1, 2, 4, 3, 9, 7, 8, 10, 14, 16]
assert min_child(minheap, 0, len(minheap)) == 1
assert min_child(minheap, 2, len(minheap)) == 5
assert min_child(minheap, 3, len(minheap)) == 7
assert min_child(minheap, 1, len(minheap)) == 3
assert min_child(minheap, 4, len(minheap)) == 9

```

<!-- grade_id: hw12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

Reorder the code for min heapify function. The function takes in the array, the index indicating which node to start the heapifying process and the size of the heap. It should modify the input array in such a way that it satisfies the min-heap property starting from the index node.

Input: 
- array: binary tree to be restored to satisfy the min-heap property
- index: index of the node where the min-heapify property should be satisfied on its subtree
- size: number of elements of the binary tree

Output:
- None, the function should modify the array in-place.

```python
from IPython.display import IFrame
IFrame('https://parsons.problemsolving.io/puzzle/be045496a109446c97b0a78e152a2f3c', 1000, 400)
```

Output:
```
<IPython.lib.display.IFrame at 0x28ae23b90>
```

Write the code below to test.

```python
def min_heapify(array: list[int|float], index: int, size: int) -> None:
   cur_idx = index
   changed = 1
   while(left_of(cur_idx) < size and changed):
      changed = 0
      min_child_idx = min_child(array, cur_idx, size)
      if (array[min_child_idx] < array[cur_idx]):
         array[min_child_idx], array[cur_idx] = array[cur_idx], array[min_child_idx]
         cur_idx = min_child_idx
         changed = 1
      
```

<!-- grade_id: hw13, points: 1, grade: true, solution: false -->
```python
array: list[int|float] = [1, 3, 4, 2, 9, 7, 8, 10, 14, 16]
min_heapify(array, 1, len(array))
assert array == [1, 2, 4, 3, 9, 7, 8, 10, 14, 16]

```

<!-- grade_id: hw14, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
def build_min_heap(array: list[int|float]) -> None:
    n = len(array)
    idx = (n//2)-1
    for i in range(idx,-1,-1):
        min_heapify(array, i, n)
```

<!-- grade_id: hw15, points: 1, grade: true, solution: false -->
```python
array: list[int|float] = [16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
print(array)
build_min_heap(array)
print(array)
assert array == [1, 2, 3, 4, 7, 9, 10, 8, 16, 14]

```

Output:
```
[16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
[1, 2, 3, 4, 7, 9, 10, 8, 16, 14]
```

<!-- grade_id: hw16, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW2.** *Heapsort:* Implement heapsort that makes use of *min-heap* instead of *max-heap*. This function returns a *new* array. The strategy is similar to max-heap, but we will use a new array to store the sorted output. Take note of the hints below:
- The top of the min-heap is always the smallest. You can take this element and put it into the output array.
- To find the next minimum, take the last element of the heap and put it into the first element of the array. Now, the tree is no longer a min-heap. Use `min_heapify()` to restore the min-heap property. This will result in a mean-heap where the first element of the array is the next minimum. You can then take out the top of the min-heap and put it into the output array.
- Reduce the heap size as you go.
- Return the new output array.

```python
import random

def gen_random_int(number: int, seed: int) -> list[int]:
    result = []
    for i in range(number+1):
        result.append(i)
    random.seed(seed)
    random.shuffle(result)
    return result
```

```python
def heapsort(array: list[int|float]) -> list[int|float]:
    result: list[int|float] = []
    build_min_heap(array)
    heap_end_pos = len(array)-1
    while(heap_end_pos > 0):
        result.append(array[0])
        array[0], array[heap_end_pos] = array[heap_end_pos], array[0]
        heap_end_pos -= 1
        min_heapify(array, 0, heap_end_pos+1)
    return result
```

<!-- grade_id: hw21, points: 1, grade: true, solution: false -->
```python
from typing import cast
array: list[int] = gen_random_int(10, 100)
array = cast(list[int|float], array)
result: list[int|float] = heapsort(array)
assert result == [0, 1, 2, 3, 4, 5, 6, 7, 8, 9]

```

<!-- grade_id: hw22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW3.** Compute the computational time for Heap Sort algorithm implemented in Python for various number of inputs. Make use of `gen_random_int(n)` to generate the input array. Use the template below to generate computation time for different number of inputs: 10, 100, 1000, etc.

```python
import time
import random
from typing import Callable

def run_function(f: Callable, x:list[int|float]) -> float:
    start: float = time.time()
    f(x)
    end: float = time.time()
    return end-start

def gen_random_int(number: int, seed: int) -> list[int]:
    result = []
    for i in range(number):
        result.append(i)
    random.seed(seed)
    random.shuffle(result)
    return result

time_heapsort: list[float] = []
# set the maximum power for 10^power number of inputs
maxpower: int = 5

for n in range(1, maxpower + 1):
    # create array for 10^1, 10^2, etc 
    # use seed 100
    array = gen_random_int((10**n),100)
    
    # call run_function with heapsort and array as arguments
    # result = run_function(None, None)
    
    result = run_function(heapsort,array)
    
    
    time_heapsort.append(result)

print(time_heapsort)
```

Output:
```
<cell>23: [1m[31merror:[m Incompatible types in assignment (expression has type [m[1m"int"[m, variable has type [m[1m"ndarray[Any, dtype[Any]]"[m)  [m[33m[assignment][m
<cell>26: [1m[31merror:[m Argument 1 to [m[1m"gen_random_int"[m has incompatible type [m[1m"ndarray[Any, dtype[signedinteger[Any]]]"[m; expected [m[1m"int"[m  [m[33m[arg-type][m
<cell>31: [1m[31merror:[m Incompatible types in assignment (expression has type [m[1m"float"[m, variable has type [m[1m"list[int | float]"[m)  [m[33m[assignment][m
<cell>31: [1m[31merror:[m Argument 2 to [m[1m"run_function"[m has incompatible type [m[1m"list[int]"[m; expected [m[1m"list[int | float]"[m  [m[33m[arg-type][m
<cell>31: [34mnote:[m [m[1m"List"[m is invariant -- see [4mhttps://mypy.readthedocs.io/en/stable/common_issues.html#variance[m[m
<cell>31: [34mnote:[m Consider using [m[1m"Sequence"[m instead, which is covariant[m
<cell>34: [1m[31merror:[m Argument 1 to [m
```

**HW4.** Plot the output of HW3 by first calculating a new x-axis computed as $n\log_2(n)$. Use the template below.

Reference:
- [Numpy Log2 function](https://docs.scipy.org/doc/numpy/reference/generated/numpy.log2.html)

```python
import matplotlib.pyplot as plt
import numpy as np

maxpower: int = 5
# create an iterable from 1 to maxpowers
powers = range(1, maxpower + 1)
# variable n stores the number of items to sort
n: list[int] = []

# Create a list of n for our x axis
for exp in powers:
    n.append(10**exp)

# convert to Numpy array
n = np.array(n)

# calculate n*log(n) for x axis 
x = n * np.log2(n)
plt.plot(x, time_heapsort, "-o")
plt.xlabel("number of input")
plt.ylabel("computation time (s)")
plt.show()
```

Output:
```
<cell>19: [1m[31merror:[m Name [m[1m"time_heapsort"[m is not defined  [m[33m[name-defined][m
NameError                                 Traceback (most recent call last)
Cell In[3], line 19
     17 # calculate n*log(n) for x axis 
     18 x = n * np.log2(n)
---> 19 plt.plot(x, time_heapsort, "-o")
     20 plt.xlabel("number of input")
     21 plt.ylabel("computation time (s)")

NameError: name 'time_heapsort' is not defined
```