# Week 4 Problem Set

## Cohort Sessions

```python
%load_ext nb_mypy
%nb_mypy On
```

Output:
```
File <frozen importlib._bootstrap>:1360, in _find_and_load(name, import_)

File <frozen importlib._bootstrap>:1324, in _find_and_load_unlocked(name, import_)

ModuleNotFoundError: No module named 'nb_mypy'
```

```python
from typing import TypeAlias
from typing import Optional, Any
import math

Number: TypeAlias = int | float
```

**CS1**. Create a class called `Square`. Its initialization takes in a variable which is the side of of a square. The object has one attribute which is called `side` to represent the side of a square. It has one property called `area` which sets the dimension of the square. Setting the area changes the side of a square.

```python
class Square:
    def __init__(self, dim: Number=1) -> None:
        self.side = dim

    @property
    def area(self) -> Number:
        return self.side * self.side
    
    @area.setter
    def area(self, value: Number) -> None:
        if value > 0:
            self.side = math.sqrt(value)

    def __str__(self) -> str:
        return f'Square with side {self.side}'


```

<!-- grade_id: cs11, points: 1, grade: true, solution: false -->
```python
s: Square = Square()
assert s.area == 1

s: Square = Square(2)
assert s.area == 4
s.area = 9
assert s.side == 3

```

<!-- grade_id: cs12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS2.** We are going to create a simple Car Racing game. First, let's create a class Car with the following properties:
- `racer` which stores the name of the driver. This property must be non-empty string. This property should be initialized upon object instantiation.
- `speed` which stores the speed of the car. This property can only be non-negative values and must be **less than or equal** to a maximum speed. The setter of this property should ensure that the value of `speed` should only falls between `0` and `max_speed` (inclusive). 
- `pos` which is an integer specifying the position of the car which can only be non-negative values.
- `is_finished` which is a computed property that returns `True` or `False` depending whether the position has reached the finish line.

Each car also has the following attributes:
- `max_speed` which specifies the maximum speed the car can have. This attribute should be initialized upon object instantiation.
- `finish` which stores the finish distance the car has to go through. Upon initialization, it should be set to -1. 

The class has the following methods:
- `start(init_speed, finish_distance)` which set the speed property to some initial value. The method also set the finish distance to some value and set the `pos` property to 0.
- `race(acceleration)` which takes in an integer value for its acceleratin and modify both the speed and the position of the car.



```python
class RacingCar:
    
    def __init__(self, name: str, max_speed: int) -> None:
        self._racer: str = name
        self.max_speed: int = max_speed
        self._speed: int = 0
        self._pos: int = 0
        self.finish: int = -1
    
    @property
    def racer(self) -> str:
        return self._racer
    
    @racer.setter
    def racer(self, name: str) -> None:
        if (name != "" and isinstance(name,str)):
            self._racer = name
            
    @property
    def speed(self) -> int:
        return self._speed
    
    @speed.setter
    def speed(self, val: int) -> None:
        if (val < 0):
            val = 0
        elif (val > self.max_speed):
            val = self.max_speed
        self._speed = val
    
    @property
    def pos(self) -> int:
        return self._pos
    
    @pos.setter
    def pos(self, val: int) -> None:
        if (val >= 0):
            self._pos = val
            
    @property
    def is_finished(self) -> bool:
        return (self.finish != -1 and self._pos >= self.finish)
            
    def start(self, init_speed: int, finish_dist: int) -> None:
        self._speed = init_speed
        self._pos = 0
        self.finish = finish_dist
    
    def race(self, acc: int) -> None:
        self._speed += acc
        self._pos += self._speed
        
    def __str__(self) -> str:
        return f"Racing Car {self.racer} at position: {self.pos}, with speed: {self.speed}."
    
            
    
```

<!-- grade_id: cs21, points: 1, grade: true, solution: false -->
```python
car: RacingCar = RacingCar("Hamilton", 200)
assert car.racer == "Hamilton"
assert car.max_speed == 200
assert car.finish == -1

car.racer = "Bottas"
assert car.racer == "Bottas"
car.racer = ""
assert car.racer == "Bottas"
car.racer = 21
assert car.racer == "Bottas"

car.speed = 10
assert car.speed == 10
car.speed = 0
assert car.speed == 0
car.speed = -10
assert car.speed == 0
car.speed = car.max_speed
assert car.speed == car.max_speed
car.speed = car.max_speed + 10
assert car.speed == car.max_speed

car.pos = 10
assert car.pos == 10
car.pos = -10
assert car.pos == 10
car.pos = 0
assert car.pos == 0

assert not car.is_finished
car.finish = 20
car.pos = 10
assert not car.is_finished
car.pos = 20
assert car.is_finished
car.pos = 22
assert car.is_finished

car.start(50, 200)
assert car.pos == 0
assert car.speed == 50
assert car.finish == 200

car.race(0)
assert car.speed == 50
assert car.pos == 50
assert not car.is_finished

car.race(10)
assert car.speed == 60
assert car.pos == 110
assert not car.is_finished

car.race(-10)
assert car.speed == 50
assert car.pos == 160
assert not car.is_finished

car.race(0)
assert car.is_finished

```

<!-- grade_id: cs22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS3**. Design a class named `Stock` to represent a company's stock that contains:
* a string computed property named `symbol` for the stock's symbol.
* a string computed property named `name` for the stock's name.
* a float property named `previous_closing_price` that stores the stock's price for the previous day.
* a float property named `current_price` that stores the stock price for the current time.
* an initializer that initializes the new stock object with the specified symbol, name, previous price and current price. 
* a method named `get_change_percent()` that returns the percentage changed from previous closing price to current price. 

```python
class Stock:
    def __init__(self, symbol: str, name: str, previous_closing_price: float, current_price: float) -> None:
        self.symbol: str = symbol
        self.name: str = name
        self.previous_closing_price: float = previous_closing_price
        self.current_price: float = current_price
        
    def get_change_percent(self):
        return (100*(self.current_price-self.previous_closing_price))/self.previous_closing_price

```

<!-- grade_id: cs31, points: 1, grade: true, solution: false -->
```python
s: Stock = Stock("INTC", "Intel Corporation", 20.5, 20.35)
assert s.symbol == "INTC"
assert s.name == "Intel Corporation"
assert s.previous_closing_price == 20.5
assert s.current_price == 20.35
assert round(s.get_change_percent(),2) == -0.73

```

**CS4.** Create a class called `Fraction` to represent a simple fraction. The class has two properties:
- `num`: which represents a numerator and of the type Integer.
- `den`: which represents a denominator and of the type Integer. Denominator should not be a zero. If a zero is assigned, you need to replace it with a `1`. 
If a `float` is assigned to either `num` or `den`, the setter should convert it to an integer using `int()` function. The setter should accept only `int` or `float` type of data.

The class should have the following method:
- `__init__(num, den)`: to initialize the numerator and the denominator. You should check if the denominator is zero. If it is you should assign `1` as the denominator instead. 
- `__str__()`:  for the object instance to be convertable to String.  You need to return a string in a format of `num/den`


```python
class Fraction:
    def __init__(self, num: int, den: int) -> None:
        self._num = int(num)
        if (den == 0):
            self._den = 1
        else:
            self._den = int(den)
    
    @property
    def num(self) -> int:
        return self._num
    
    @num.setter
    def num(self, val: int) -> None:
        self._num = int(val)
    
    @property
    def den(self) -> int:
        return self._den
    
    @den.setter
    def den(self, val: int) -> None:
        self._den = int(val)
    
    def __str__(self) -> str:
        ret = str(self._num) + "/" + str(self._den)
        return ret
    
```

<!-- grade_id: cs41, points: 1, grade: true, solution: false -->
```python
f0: Fraction = Fraction(0, 1)
assert f0.num == 0
assert f0.den == 1
assert str(f0) == "0/1"

f1: Fraction = Fraction(1, 2)
assert f1.num == 1
assert f1.den == 2
assert str(f1) == "1/2"

f1.num = 3
f1.den = 4
assert str(f1) == "3/4"

```

<!-- grade_id: cs42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS5.** Implement the `Stack` abstract data type using a Class. You can use `list` Python data type as its internal data structure. Name this `list` as `items`.

The class should have the following interface:
- `__init__()` to initialize an empty list for the stack to store the items.
- `push(item)` which stores an item into the top of the stack.
- `pop()` which returns and removes the top element of the stack. The return value is optional as it may return `None` if there are no more elements in the stack.
- `peek()` which returns the top element of the stack. If the stack is empty, it returns `None`.

The class should have the following properties:
- `is_empty` is a computed property which returns either `True` or `False` depending whether the stack is empty or not.
- `size` is a computed property which returns the number of items in the stack.


```python
class Stack:
    def __init__(self) -> None:
        self.__items: list[Any] = []
        
    def push(self, item: Any):
        self.__items.append(item)

    def pop(self) -> Any:
        if (len(self.__items) == 0):
            return None
        return self.__items.pop()

    def peek(self) -> Any:
        return self.__items[-1]

    @property
    def is_empty(self) -> bool:
        return (len(self.__items) == 0)

    @property
    def size(self):
        return len(self.__items)


```

<!-- grade_id: cs51, points: 1, grade: true, solution: false -->
```python
s1: Stack = Stack()
s1.push(2)
assert not s1.is_empty
assert s1.pop() == 2
assert s1.is_empty
assert s1.pop() == None
s1.push(1)
s1.push(2)
s1.push(3)
assert not s1.is_empty
assert s1._Stack__items == [1, 2, 3]
assert s1.peek() == 3
assert s1.size == 3

```

<!-- grade_id: cs52, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS6.** Implement a Queue abstract data structure using two Stacks instead of Python's list. For this double-stack implementation, the Queue has a *left* Stack and a *right* Stack. The enqueue and dequeue operations work as follows:
- enqueue: This operation just pushes the new item into the *left* Stack.
- dequeue: This operation is done as follows:
    - if the *right* Stack is empty: create a new *right* Stack which is the reverse of the items in the *left* Stack. You should then empty the *left* stack.
    - if the *right* Stack is not empty: pop from the *right* Stack.
- peek: This operation is similar to Stack's peek. It returns the value at the front of the Queue without removing it. If the Queue is empth, it should return a `None` object.

Note that, you could also enqueue to the right Stack and dequeue from the left stack. You can choose to implement according to this if you wish and make sure the operations still work as defined.

```python
class Queue:
    def __init__(self) -> None:
        self.left_stack: Stack = Stack()
        self.right_stack: Stack = Stack()
        
    def enqueue(self, item: Any):
        self.left_stack.push(item)

    def dequeue(self) -> Any:
        if (self.right_stack.is_empty):
            while(self.left_stack.size > 0):
                self.right_stack.push(self.left_stack.pop())
        return self.right_stack.pop()
            

    def peek(self) -> Any:
        if (self.right_stack.size + self.left_stack.size == 0):
            return None
        if (self.right_stack.size > 0):
            return self.right_stack.peek()
        while(self.left_stack.size > 0):
            self.right_stack.push(self.left_stack.pop())
        return self.right_stack.peek()

    @property
    def is_empty(self) -> bool:
        return (self.left_stack.size + self.right_stack.size == 0)

    @property
    def size(self):
        return self.left_stack.size + self.right_stack.size

```

<!-- grade_id: cs61, points: 1, grade: true, solution: false -->
```python
q1: Queue = Queue()
q1.enqueue(2)
assert not q1.is_empty 
assert q1.size == 1
assert q1.dequeue() == 2
assert q1.is_empty
q1.enqueue(1)
q1.enqueue(2)
q1.enqueue(3)
assert q1.size == 3
assert q1.peek() == 1
assert q1.dequeue() == 1
assert q1.dequeue() == 2
assert q1.dequeue() == 3
assert q1.peek() == None

```

<!-- grade_id: cs62, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```