# Week 6 Problem Set

## Cohort Sessions

```python
%load_ext nb_mypy
%nb_mypy On
```

Output:
```
File <frozen importlib._bootstrap>:1360, in _find_and_load(name, import_)

File <frozen importlib._bootstrap>:1324, in _find_and_load_unlocked(name, import_)

ModuleNotFoundError: No module named 'nb_mypy'
```

```python
from typing import TypeAlias
from typing import Optional, Any, Callable, Iterator, Iterable, cast
from __future__ import annotations

Number: TypeAlias = int | float
```

**CS1.** *Dictionary:* Implement a Graph using a *Dictionary* where the keys are the Vertices in the Graph and the values (in the the key-value pair) correspond to an Array containing the neighbouring Vertices. For example, let's represent the following graph:
```
    A -> B
    A -> C
    B -> C
    B -> D
    C -> D
    D -> C
    E -> F
    F -> C
```
Create a Dictionary to represent the graph above.

```python
# replace the None with a dictionary representing the graph
graph: Optional[dict[str, list[str]]] = {"A": ["B", "C"], 
         "B": ["C", "D"],
         "C": ["D"],
         "D": ["C"], 
         "E": ["F"],
         "F": ["C"]}
```

<!-- grade_id: cs1, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS2.** Create a function `create_bfs_tree(graph, start, end)`. The function takes in three input arguments. The first one is the graph which is represented as a dictionary in adjacency list described in the earlier cohort questions. The second one is the label of the starting vertex to begin the Breadth-First Search. The third argument is the label of the ending vertex. 

The function should return a dictionary containing the parent node of each vertex in the path from start to end. For example, Let's say the tree generated is shown below.

```
A
| \
B  D
|  |
C  E
|  |
F
```

The output dictionary should be something like the following.

```
{'F': 'C', 'C': 'B', 'E':'D', 'B':'A', 'D':'A'}
```

Note: Feel free to make use of the `get_neighbours()` function from your HW or your own code to get the neighbours from the values of the dictionary.

```python
def get_neighbours(graph: dict[str, list[str]], vert: str) -> Optional[list[str]]:
    return graph[vert]
```

```python
def create_bfs_tree(graph: dict[str, list[str]], start: str, end: str) -> dict[str, Optional[str]]:
    output: dict[str, Optional[str]] = {'A': None}
    visited = [start]
    to_explore = [start]
    
    while (len(to_explore)):
        topz = to_explore.pop(0)
        neighbours = get_neighbours(graph, topz)
        if (neighbours is None):
            continue
        for neigh in neighbours:
            if (neigh in visited):
                continue
            visited.append(neigh)
            to_explore.append(neigh)
            output[neigh] = topz
    
    return output
```

<!-- grade_id: cs21, points: 1, grade: true, solution: false -->
```python
graph: dict[str, list[str]] = {"A": ["B", "D"],
         "B": ["A", "C"],
         "C": ["B", "D", "F"],
         "D": ["A", "C", "E"],
         "E": ["D", "F"],
         "F": ["C", "E"]}

output: dict[str, Optional[str]] = create_bfs_tree(graph, "A", "F")
print(output)
assert output == {'A': None, 'B': 'A', 'D': 'A', 'C': 'B', 'E': 'D', 'F': 'C'}
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

Output:
```
{'A': None, 'B': 'A', 'D': 'A', 'C': 'B', 'E': 'D', 'F': 'C'}
```

<!-- grade_id: cs22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS3.** Create a subclass of `Vertex` called `SearchVertex`. This class has the following additional attributes:
- `colour`: which is a mark on the vertex during the search algorithm. It is of String data type and should be set to "white" by default.
- `d`: which is an Integer denoting the distance from other Vertex to the current Vertex in Breath-First-Search. This is also used to record discovery time in Depth-First-Search. This attribute should be initialized to `sys.maxsize` at the start.
- `f`: which is an Integer denoting the final time in Depth-First-Search. This attribute should be initialized to `sys.maxsize` at the start.
- `parent`: which is a reference to the parent Vertex object. This attribute should be set to `None` at the start.

```python
class Vertex:
    def __init__(self, id_: str="") -> None:
        self.id_: str = id_
        self.neighbours: dict[Vertex, Number] = {}
    
    def add_neighbour(self, nbr_vertex: Vertex, weight: Number=0) -> None:
        self.neighbours[nbr_vertex] = weight
    
    def get_neighbours(self) -> list[Vertex]:
        return list(self.neighbours)
    
    def get_weight(self, neighbour: Vertex) -> Optional[Number]:
        return self.neighbours.get(neighbour, None)
        #return self.neighbours[neighbour]
    
    def __eq__(self, other) -> bool:
        return self.id_ == other.id_
    
    def __lt__(self, other) -> bool:
        return self.id_ < other.id_
    
    def __hash__(self) -> int:
        return hash(self.id_)
    
    def __str__(self) -> str:
        ret = "Vertex " + str(self.id_) + " is connected to:"
        tek = 0
        for vert in self.neighbours:
            if (tek == 0):
                tek = 1
                ret += " " + str(vert.id_)
            else:
                ret += ", " + str(vert.id_)
        #print(ret)
        return ret
```

```python
import sys

class SearchVertex(Vertex):
    def __init__(self, id_: str="") -> None:
        super().__init__(id_)
        self.colour: str = "white"
        self.d: int = sys.maxsize
        self.f: int = sys.maxsize
        self.parent: Vertex = None

```

<!-- grade_id: cs31, points: 1, grade: true, solution: false -->
```python
import sys

v:SearchVertex = SearchVertex()
assert v.id_ == ""
assert v.colour == "white"
assert v.d == sys.maxsize
assert v.f == sys.maxsize
assert v.parent == None

parent_method: Any = getattr(v, 'get_neighbours', None)
assert callable(parent_method)
parent_method: Any = getattr(v, 'get_weight', None)
assert callable(parent_method)

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: cs32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS4.** Create a class called `SearchGraph` which is a subclass of the class `Graph`. This class should override the method `_create_vertex(id)` and instantiate a `SearchVertex` object instead of `Vertex`.

```python
class Graph:
    def __init__(self) -> None:
        self.vertices: dict[str, Vertex] = {}
        
    def _create_vertex(self, id_: str) -> Vertex:
        return Vertex(id_)
    
    def add_vertex(self, id_: str) -> None:
        self.vertices[id_] = self._create_vertex(id_)
    
    def get_vertex(self, id_: str) -> Optional[Vertex]:
        if (id_ in self.vertices):
            return self.vertices[id_]
        return None
    
    def add_edge(self, start_v: str, end_v: str, weight: Number=0) -> None:
        if (self.get_vertex(start_v) is None):
            self.add_vertex(start_v)
        
        if (self.get_vertex(end_v) is None):
            self.add_vertex(end_v)
        
        self.vertices[start_v].add_neighbour(self.get_vertex(end_v))
        
    def get_neighbours(self, id_: str) -> list[str]:
        if (id_ in self.vertices):
            ret = []
            for vert in self.vertices[id_].neighbours:
                ret.append(vert.id_)
            return ret
        else:
            return []
    
    def __contains__(self, val: str) -> bool:
        return val in self.vertices.keys()
    
    def __iter__(self):
        for k,v in self.vertices.items():
            yield v 
        
    @property
    def num_vertices(self):
        return int(len(self.vertices))
```

```python
class SearchGraph(Graph):
    def _create_vertex(self, id_):
        return SearchVertex(id_)
```

<!-- grade_id: cs41, points: 1, grade: true, solution: false -->
```python
g2: SearchGraph = SearchGraph()
g2.add_vertex("Z")
assert(type(g2.vertices["Z"]) == type(SearchVertex()))
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: cs42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS5.** Create a class `TraverseGraph` which takes in an object `SearchGraph` for its initialization. The class should have the following methods:
- `clear_vertices()`: which sets the attributes f all the vertices:
  - `colour` to "white"
  - `d` to `sys.maxsize`
  - `f` to `sys.maxsize`
  - `parent` to `None`.
 




```python
import sys

class TraverseGraph:
    def __init__(self, g: SearchGraph) -> None:
        self.graph = g
    
    def clear_vertices(self) -> None:
        for vert in self.graph.vertices.values():
            vert.colour = "white"
            vert.d = sys.maxsize
            vert.f = sys.maxsize
            vert.parent = None
    
    def __iter__(self) -> Iterator:
        return iter([v for v in self.graph])
    
    def __len__(self) -> int:
        return len([v for v in self.graph.vertices])
```

<!-- grade_id: cs51, points: 1, grade: true, solution: false -->
```python
g4: SearchGraph = SearchGraph()
g4.add_vertex("A")
g4.add_vertex("B")
g4.add_vertex("C")
g4.add_vertex("D")
g4.add_vertex("E")
g4.add_vertex("F")
g4.add_edge("A", "B")
g4.add_edge("A", "C")
g4.add_edge("B", "C")
g4.add_edge("B", "D")
g4.add_edge("C", "D")
g4.add_edge("D", "C")
g4.add_edge("E", "F")
g4.add_edge("F", "C")
gs4: TraverseGraph = TraverseGraph(g4)
gs4.clear_vertices()

assert len(gs4) == 6
assert [v.id_ for v in gs4] == ["A", "B", "C", "D", "E", "F"]
assert [v.colour for v in gs4] == ["white" for v in range(len(gs4))]
assert [v.d for v in gs4] == [sys.maxsize for v in range(len(gs4))]
assert [v.f for v in gs4] == [sys.maxsize for v in range(len(gs4))]
assert [v.parent for v in gs4] == [None for v in range(len(gs4))]
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: cs52, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS6.** Create a class `TraverseBFS` which is a subclass of `TraverseGraph`. This subclass should implement the Breadth First Search algorithm in the following methods:

- `search_from(start)`: which initializes the `d` and `parent` attributes of each vertices in the graph from the `start` Vertex following Breadth-First-Search algorithm. Use your previous code that implements `Queue` data structure. 
- `get_shortest_path(start, dest)`: which returns a list of vertex ids that forms a shortest path from Vertex `start` to Vertex `dest`. This method should call `get_path()` (see next method in the list) and pass on an empty list as one of the input arguments. The method `get_path()` will populate this list if there is a path. 
    - If the path list is empty after calling `get_path()`, this means that either the starting vertex or the destination vertex do not exist in the grapth. In this case, exit the function returning a `None` object.
    - If the path list is not empty, it will either contain `No Path` as one of the items or a list of vertices that gives the path from the starting vertex to the destination vertex. In this case, simply return the list as it is. 
- `get_path(start, dest, result)`: which modifies the input list `result`. 
    - This method should first check whether the starting vertex and the destination vertex exist in the grapth. If they do not exist in either case, the method should exit returning a `None` object. 
    - If the starting and destination vertex exist in the graph, this method should call `search_from()` when the distance at `start` Vertex is not zero. A non-zero distance at the starting vertex means that we have not run the BFS algorithm from that starting vertex.
    - if the destination and the starting vertex are the same, modify the `result` list with this one vertex. This means that we have found the path that consists of only one vertex. 
    - if the destination vertex has no parent, this means there is no path. Add `No Path` string into the `result` list. 
    - otherwise, recursively call `get_path()` and add the result into the `result` list. 

```python
class Stack:
    def __init__(self) -> None:
        self.__items: list[Any] = []
        
    def push(self, item: Any):
        self.__items.append(item)

    def pop(self) -> Any:
        if (len(self.__items) == 0):
            return None
        return self.__items.pop()

    def peek(self) -> Any:
        return self.__items[-1]

    @property
    def is_empty(self) -> bool:
        return (len(self.__items) == 0)

    @property
    def size(self):
        return len(self.__items)



class Queue:
    def __init__(self) -> None:
        self.left_stack: Stack = Stack()
        self.right_stack: Stack = Stack()
        
    def enqueue(self, item: Any):
        self.left_stack.push(item)

    def dequeue(self) -> Any:
        if (self.right_stack.is_empty):
            while(self.left_stack.size > 0):
                self.right_stack.push(self.left_stack.pop())
        return self.right_stack.pop()
            

    def peek(self) -> Any:
        if (self.right_stack.size + self.left_stack.size == 0):
            return None
        if (self.right_stack.size > 0):
            return self.right_stack.peek()
        while(self.left_stack.size > 0):
            self.right_stack.push(self.left_stack.pop())
        return self.right_stack.peek()

    @property
    def is_empty(self) -> bool:
        return (self.left_stack.size + self.right_stack.size == 0)

    @property
    def size(self):
        return self.left_stack.size + self.right_stack.size
```

```python
class TraverseBFS(TraverseGraph):

    def search_from(self, start: str) -> None:
        self.clear_vertices()
        start_v = self.graph.get_vertex(start)
        if (start_v is None):
            return None

        start_v.colour, start_v.d, start_v.parent = "grey", 0, None
        queue = Queue()
        queue.enqueue(start_v)
        
        while queue.size:
            top_v = queue.dequeue()
            for nx in top_v.get_neighbours():
                if (nx.colour == "white"):
                    queue.enqueue(nx)
                    nx.d = top_v.d + 1
                    nx.colour = "grey"
                    nx.parent = top_v
            top_v.colour = "black"
            
    
    def get_shortest_path(self, start: str, dest: str) -> Optional[list[str]]:
        ret = []
        self.get_path(start, dest, ret)
        
        if (ret == []):
            return None
        #print(ret)
        return ret
    
    def get_path(self, start: str, dest: str, result: list[str]) -> None:
        if (self.graph.get_vertex(start) is None):
            return None
        if (self.graph.get_vertex(dest) is None):
            return None

        start_v = self.graph.get_vertex(start)
        dest_v = self.graph.get_vertex(dest)
        if (start_v.d != 0):
            self.search_from(start)
        
        if (start == dest):
            result.append(start)
        elif (dest_v.parent is None):
            result.append("No Path")
        else:
            self.get_path(start, dest_v.parent.id_, result)
            result.append(dest)

```

<!-- grade_id: cs61, points: 1, grade: true, solution: false -->
```python
g4: SearchGraph = SearchGraph()
g4.add_vertex("A")
g4.add_vertex("B")
g4.add_vertex("C")
g4.add_vertex("D")
g4.add_vertex("E")
g4.add_vertex("F")
g4.add_edge("A", "B")
g4.add_edge("A", "C")
g4.add_edge("B", "C")
g4.add_edge("B", "D")
g4.add_edge("C", "D")
g4.add_edge("D", "C")
g4.add_edge("E", "F")
g4.add_edge("F", "C")
gs4: TraverseBFS = TraverseBFS(g4)

gs4.search_from("A")
assert cast(SearchVertex, gs4.graph.get_vertex("A")).d == 0
assert cast(SearchVertex, gs4.graph.get_vertex("A")).colour == "black"
assert cast(SearchVertex, gs4.graph.get_vertex("A")).parent == None
assert cast(SearchVertex, gs4.graph.get_vertex("B")).d == 1
assert cast(SearchVertex, gs4.graph.get_vertex("B")).colour == "black"
assert cast(SearchVertex, gs4.graph.get_vertex("B")).parent == gs4.graph.get_vertex("A")
assert cast(SearchVertex, gs4.graph.get_vertex("C")).d == 1
assert cast(SearchVertex, gs4.graph.get_vertex("C")).colour == "black"
assert cast(SearchVertex, gs4.graph.get_vertex("C")).parent == gs4.graph.get_vertex("A")
assert cast(SearchVertex, gs4.graph.get_vertex("D")).d == 2
assert cast(SearchVertex, gs4.graph.get_vertex("D")).colour == "black"
cast(SearchVertex, gs4.graph.get_vertex("D")).parent
#assert gs4.graph.get_vertex("D").parent == gs4.graph.get_vertex("B")
ans: Optional[list[str]] = gs4.get_shortest_path("A", "D")
assert ans == ["A", "B", "D"]
ans: Optional[list[str]] = gs4.get_shortest_path("E", "D")
assert ans == ["E", "F", "C", "D"]
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: None, points: 0, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```