# Week 08 Problem Set

## Cohort Sessions

```python
%load_ext nb_mypy
%nb_mypy On
```

Output:
```
File <frozen importlib._bootstrap>:1360, in _find_and_load(name, import_)

File <frozen importlib._bootstrap>:1324, in _find_and_load_unlocked(name, import_)

ModuleNotFoundError: No module named 'nb_mypy'
```

```python
from typing import TypeAlias
from typing import Optional, Any, Callable, Iterator, Iterable, cast
from __future__ import annotations

Number: TypeAlias = int | float
```

**CS1.** Define an Abstract Class for a State Machine, called `StateMachine`. The class has one attribute and one computed property:
- `state`: which is the current state of the machine and is the attribute of any `StateMachine` object instance.
- `start_state`: which is the initial state of the machine and is the computed property to be defined in the child class.

The class should define the following methods:
- `start()`: this method set the `state` property using the value in `start_state`. Once `state` has a value, the machine is considered started.
- `step(inp)`: this method takes in the current input and returns the current output. This method should move the state machine to the next state based on the current input and its current state. You should call `get_next_values(state, inp)` in your implementation.
- `done(state)`: this method always return `False`. A child class can override thid method to give a different condition to end the state machine.
- `is_done()`: is to be used internally to check if the state machine should terminate or not. This method simply calls `done(state)` and pass on the current `state`. The method `transduce(inp_list)` calls this method to check if it should terminates or not.
- `transduce(inp_list)`: this method calls `start()` to initialize the `state` with the `start_state` and run the state machine by calling `step(inp)` for every item in the `inp_list`. The method runs the state machine and produces the output list according to the number of input in the `inp_list` or when the state machine terminates according to the output of `is_done()` method. This method should call `is_done()` to see if it should terminate at a particular state.

This class should be an Abstract Class. Implement the following way:
- `SM` class inherits from `abc.ABC`, which is Python's Abstract Base Class (ABC). 
- Any implementation of State Machine's instances must declare the following property: `start_state`.
- Any implementation of State Machine's instances must implement the following abstract method: `get_next_values()` that takes in the current `state` and the the current `input` and output a tuple of the `next_state` and the current `output` . 


```python
from abc import ABC, abstractmethod

class StateMachine(ABC):
    
    def start(self):
        self.state = self.start_state

    def step(self, inp):
        prev = self.state
        cur_state, output = self.get_next_values(prev,inp)
        self.state = cur_state
        return output
        
    def transduce(self, inp_list):
        self.start()
        ret = []
        for item in inp_list:
            if (self.is_done()):
                return ret
            ret.append(self.step(item))
        return ret
        

    @property
    @abstractmethod
    def start_state(self):
        pass 

    @abstractmethod
    def get_next_values(self, state, inp):
        pass

    def done(self, state):
        return False

    def is_done(self):
        return self.done(self.state)

    
```

<!-- grade_id: cs11, points: 1, grade: true, solution: false -->
```python
class Test(StateMachine):
    @property
    def start_state(self) -> int:
        return 0

    def get_next_values(self, state: int, inp: int) -> tuple[int, int]:
        next_state = state + inp
        output = next_state
        return next_state, output

    def done(self, state: int) -> bool:
        if state == -1:
            return True
        else:
            return False

class NoImplement(StateMachine):
    pass
    
t1 = Test()
t1.start()
assert t1.state == 0
out = t1.step(2)
assert t1.state ==2 and out == 2

t2 = Test()
out = t2.transduce([1,2,3,4])
assert out == [1, 3, 6, 10]

t3 = Test()
out = t3.transduce([1, -2, 3])
assert out == [1, -1]

try:
    t4 = NoImplement()
    raise AssertionError
except TypeError:
    pass

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: cs12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS2.** *Coke Machine:* In this problem, you will implement in Python the behavior of a simplified coke-dispensing machine. The behavior of such a machine is captured in the state diagram shown in the Figure below. The machine consists of two states labelled 0 and 1. The state diagram does not show what the machine would do if an unexpected coin is inserted. Therefore, assume that any unexpected coin is returned to the user without a change in the machines state. Thus, on your own, you may want to complete the Figure below to add in the missing transitions.

![](https://www.dropbox.com/s/kzk6nkdss7wvw85/coke_sm.png?raw=1)

Each directed arc in the state diagram is labelled as $x/y$ where $x$ denotes the input received and $y$, the output generated. For example, the arc that connects state 0 to state 1 that’s labelled `50/(50, ’--’,0)` means that when the dispenser receives 50¢ (50 before the /) in state 0, it moves to state 1 and generates an output of `(50, ’--’,0)`. This tuple of values in the output indicates that the dispenser display shows 50 which is the amount deposited by the user, no coke has been dispensed yet as indicated by `--`, and no change has been returned to the user as indicated by the last entry which is a 0.
The machine accepts only 50¢ and one dollar (100¢) coins. It has a display that shows how many cents have been deposited.
- State 0: When a 50¢ coin is deposited the dispenser moves to state 1. At this moment in time, the display shows 50 but nothing is dispensed and no change is returned. If a dollar coin is deposited, the machine continues to display 0, dispenses coke, and does not return any money (well, why should it!).
- State 1: When a 50¢ coin is deposited the dispenser moves to state 0. At this moment in time, the display shows 0, coke is dispensed and no change is returned. If a dollar coin is deposited the machine continues to display 0, dispenses coke, and returns 50¢.

We wish to write a program that simulates the behavior of the coke dispenser as described above. We will write a class named CokeMachine that contains properties and methods as described below:
- `CokeMachine` class is a subclass of `StateMachine` class.
- `CokeMachine` class has a class attribute called start_state which is the starting state of the machine. This attribute should be initialized to 0, which represents state 0 in the diagram above.
- `CokeMachine` class implements the abstract method `get_next_values(state, inp)` that takes in the current state and the input, and returns the next state and output as a tuple. Think about the following: which state represents the following scenarios?
    - the coke machine is waiting for a valid coin to be deposited 
    - the coke machine has a 50-cent coin in it

Sample Interaction:
```
cm = CokeMachine()
cm.start()
print(c.step(50))
print(c.step(50))
print(c.step(100))
print(c.step(10))
print(c.step(50))
print(c.step(100))
print(c.step(10))
```

The output should be:
```
(50, "--", 0)
(0, "coke", 0)
(0, "coke", 0)
(0, "--", 10)
(50, "--", 0)
(0, "coke", 50)
(0, "--", 10)
``` 

```python
class CokeMachine(StateMachine):

    @property
    def start_state(self) -> int:
       self.state = 0
       return self.state

    def get_next_values(self, state: int, inp: int) -> tuple[int, tuple[int, str, int]]:
        next_state, output = 0,0
        if (state == 0):
            if (inp == 50):
                next_state = 1
                output = (50, "--", 0)
            elif (inp == 100):
                next_state = 0
                output = (0, "coke", 0)
            else:
                output = (0, "--", inp)
        else:
            if (inp == 50):
                next_state = 0
                output = (0, "coke", 0)
            elif (inp == 100):
                next_state = 0
                output = (0, "coke", 50)
            else:
                output = (0, "--", inp)
            
        return next_state, output
```

<!-- grade_id: cs21, points: 1, grade: true, solution: false -->
```python
cm: CokeMachine = CokeMachine()
cm.start()
assert cm.state == 0
out: tuple[int, str, int] = cm.step(50)
assert out == (50, "--", 0) and cm.state == 1
out: list[int] = cm.transduce([50, 50, 100, 10, 50, 100, 10])
assert out == [(50, '--', 0), (0, 'coke', 0), (0, 'coke', 0), (0, '--', 10), (50, '--', 0), (0, 'coke', 50), (0, '--', 10)]

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: cs22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS3.** *Simple Account:* In this problem, you will need to create a state machine that simulates a simple bank account. This is similar to the Accumulator state machine in the notes. The only difference is that any withdrawal when the balance is less than \\$100  incurs a \\$5 charge. The state machine should fulfill the following:
- The starting balance is specified when instantiating the object.
- The output of the state machine is the current balance after the transaction.

Sample interaction:
```
acct = SimpleAccount(110)
acct.start()
print(acct.step(10))
print(acct.step(-25))
print(acct.step(-10))
print(acct.step(-5))
print(acct.step(20))
print(acct.step(20))
```

The expected output is:
```
120
95
80
70
90
110
```

In the code template below, we have provided a general getter and setter implementation for the computed property `start_state`. In this task, you need to set this computed property in the initializer `__init__()`.

```python
class SimpleAccount(StateMachine):
    @property
    def start_state(self) -> Number:
        return self.__start_state
    
    @start_state.setter
    def start_state(self, value: Number) -> None:
        self.__start_state: Number = value

    def __init__(self, balance: Number) -> None:
        self.start_state = balance

    def get_next_values(self, state: Number, inp: Number) -> tuple[Number, Number]:
        next_state, output = 0, 0
        if (inp >= 0):
            next_state = state + inp
        else:
            if (state < 100):
                inp -= 5
            next_state = state + inp
        
        output = next_state
        return next_state, output
```

<!-- grade_id: cs31, points: 1, grade: true, solution: false -->
```python
acct: SimpleAccount = SimpleAccount(110)
out: list[Number] = acct.transduce([10, -25, -10, -5, 20, 20])
assert out == [120, 95, 80, 70, 90, 110]

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: cs32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS4.** *Graph:* We can consider graph search problem as a state-space search, where each state is a node in the graph and the transition between one state to another as an edge in the graph. Recall that we can represent graph using either a Dictionary or a Class. In this problem, we will use Dictionary to represent a graph where each node represents a state. For example, the image below can be represented as follows:

```
map = {"S": ["A", "B"],
        "A": ["S", "C", "D"],
        "B": ["S", "D", "E"],
        "C": ["A", "F"],
        "D": ["A", "B", "F", "H"],
        "E": ["B", "H"],
        "F": ["C", "D", "G"],
        "H": ["D", "E", "G"],
        "G": ["F", "H"]}
```

<img src="https://data-driven-world.github.io/2023/assets/images/state_space_map-11c07cbe83c95b6f8f8e9915f832ee3e.png" width="600"></img>

Let's consider an action to be the index of the list in that Dictionary. For example, if my current state is "S", action 0 will give "A" as the next state while action 1 will give "B" as the next state. 

Write `MapSM` class and override the `get_next_values()` method to 
- `__init__(start)`: which initialize the `start_state` with `start` during object instantiation.
- `get_next_values(state, inp)`: this method produces the next state based on the input and the current state. The next states are the neighbours of the current vertex based on the `inp`. The `inp` argument is the index of the next vertex to be visited. The state machine should remain in the current state if the input is not valid or when the current vertex has no neighbours.

The class `MapSM` should implement `StateSpaceSearch` class which is an abstract class inheriting from `StateMachine` abstract class. The `StateSpaceSearch` abstract class requires `MapSM` to implement the properties:
- `statemap`: which gives the dictionary or the map of the state space.
- `legal_inputs`: which gives a `set` of legal inputs of this state space. This is a set of integers from 0 up to $n-1$ where $n$ is the largest number of neighbours of a state in the graph. 


```python
from abc import abstractmethod

class StateSpaceSearch(StateMachine):
    
    @property
    @abstractmethod
    def statemap(self):
        pass

    @property
    @abstractmethod
    def legal_inputs(self):
        pass

    @property
    def start_state(self):
        return self.__start_state
    
    @start_state.setter
    def start_state(self, value):
        self.__start_state = value
    
```

```python
class MapSM(StateSpaceSearch):
        
    def __init__(self, start: str) -> None:
        self.start_state = start

    @property
    def statemap(self) -> dict[str, list[str]]:
        statemap = {"S": ["A", "B"],
                    "A": ["S", "C", "D"],
                    "B": ["S", "D", "E"],
                    "C": ["A", "F"],
                    "D": ["A", "B", "F", "H"],
                    "E": ["B", "H"],
                    "F": ["C", "D", "G"],
                    "H": ["D", "E", "G"],
                    "G": ["F", "H"]}
        return statemap

    @property
    def legal_inputs(self) -> set[int]:
        n = 0
        for key, value in self.statemap.items():
            n = max(n, len(value))
        
        return set(range(0,n))

  
    def get_next_values(self, state: str, inp: int) -> tuple[str, str]:
        neigh = self.statemap[state]
        next_state = state
        if (len(neigh) > inp):
            next_state = neigh[inp]
        return next_state, next_state
            
        
            
    
```

<!-- grade_id: cs41, points: 1, grade: true, solution: false -->
```python
mapSM: MapSM = MapSM("S")
mapSM.start()
ans: list[int] = mapSM.transduce([0, 1, 1, 2, 0])
assert ans == ["A", "C", "F", "G", "F"]
assert mapSM.legal_inputs == set(range(4))

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: cs42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS5.** Create a class to represent a node in a state-space search called `SearchNode`. The object instance of `SearchNode` has the following properties which should be initialized during instantiation.
- `state`: which is the state or the label of the node.
- `action`: which is the action that was taken to arrive at the node.
- `parent`: which is the parent search node from which the current search node can be reached. If a node has no parent, this property should be initialized to `None`.

The class has the following method:
- `path()`: which returns a list of `Step` from the root node to the current node. The `Step` object is a class and is defined in the template. A `Step` object has two properties: `action` and `state`. You should use recursion for this method `path()`. The base case is when the node has no parent. In this case the solution contains only one step which is the action and the state of the current node. The recursive case is when the current has a `parent` object. In this case, you have to traverse to the ancestors' node until it reach a node which has no `parent` object. 
- `in_path(state)`: which returns `True` if the state in the argument is in the path. *Hint: use recursion and the parent's `in_path()` method.*

```python
class Step:
    def __init__(self, action: Optional[int], state: str) -> None:
        self.action: Optional[int] = action
        self.state: str = state
    
    def __eq__(self, other) -> bool:
        return self.action == other.action and self.state == other.state
  
    def __str__(self) -> str:
        return f"action: {self.action:}, state: {self.state:}"
```

```python
class SearchNode:
    def __init__(self, action: Optional[int], state: str, parent: Optional[SearchNode]) -> None:
        self.state: str = state
        self.action: Optional[int] = action
        self.parent: Optional[SearchNode] = parent
  
    def path(self) -> list[Step]:
        if (self.parent == None):
            return [Step(self.action, self.state)]
        return self.parent.path() + [Step(self.action, self.state)]
  
    def in_path(self, state: str) -> bool:
        if self.state == state:
            return True
        elif self.parent is None:
            return False
        else:
            return self.parent.in_path(state)
  
    def __eq__(self, other) -> bool:
        if self is None and other is None:
            return True
        elif self is None:
            return False
        elif other is None:
            return False
        else:
            return self.state == other.state and self.parent == other.parent and \
                   self.action == other.action
```

<!-- grade_id: cs51, points: 1, grade: true, solution: false -->
```python
s: SearchNode = SearchNode(None, "S", None)
a:  SearchNode = SearchNode(0, "A", s)
b: SearchNode  = SearchNode(1, "B", s)
s1: SearchNode  = SearchNode(0, "S", a)
c: SearchNode  = SearchNode(1, "C", a)
d1: SearchNode  = SearchNode(2, "D", a)
s2: SearchNode  = SearchNode(0, "S", b)
d2: SearchNode  = SearchNode(1, "D", b)
e: SearchNode = SearchNode(2, "E", b)
a1: SearchNode  = SearchNode(0, "A", s1)
b1: SearchNode  = SearchNode(1, "B", s1)
a2: SearchNode  = SearchNode(0, "A", c)
f1: SearchNode  = SearchNode(1, "F", c)
a3: SearchNode  = SearchNode(0, "A", d1)
b2: SearchNode  = SearchNode(1, "B", d1)
f2: SearchNode  = SearchNode(2, "F", d1)
h1: SearchNode  = SearchNode(3, "H", d1)
a4: SearchNode  = SearchNode(0, "A", s2)
b3: SearchNode  = SearchNode(1, "B", s2)
a5: SearchNode  = SearchNode(0, "A", d2)
b4: SearchNode  = SearchNode(1, "B", d2)
f3: SearchNode  = SearchNode(2, "F", d2)
h2: SearchNode  = SearchNode(3, "H", d2)
b5: SearchNode  = SearchNode(0, "B", e)
h3: SearchNode  = SearchNode(1, "H", e)

assert s.parent == None
assert a.state == "A" and a.parent == s and a.action == 0
assert b.state == "B" and b.parent == s and b.action == 1
assert h3.state == "H" and h3.parent == e and h3.action == 1
assert a5.path() == [Step(None, "S"), Step(1, "B"), Step(1, "D"), Step(0, "A")]
assert b5.in_path("B")
assert b5.in_path("S")
assert b5.in_path("E")

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: cs52, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```